<a href="https://colab.research.google.com/github/gkrivitsky/CS1010-LAB/blob/main/ps3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Problem Set 3
#### CMSI 1010: Computer Programming and Lab

# Important Dates

 - Released: 2024-02-08 at 12:00 am [PT]
 - Deadline: 2024-02-22 at 11:59 pm [PT]

# Question 1: Reading Notes

All readings are from the [_Think Python_ textbook](http://greenteapress.com/thinkpython2/thinkpython2.pdf):
* Supporting content for this pset: Chapters 7–9
* Upcoming content: Chapters 11–13, 15–17

In your own words, summarize the key topics covered in the textbook chapters. Additionally, relate these topics to the concepts and material we've learned in class. How do the ideas presented in the textbook connect to the lessons and discussions we've had in the classroom? Please provide specific examples to support your points.

Functions and Scope: Chapters 7 emphasizes the importance of functions in structuring code and managing scope. This ties into our classroom discussions about how functions help avoid repetition and improve code readability. For instance, when we created modular code for our projects, we relied heavily on defining clear functions to encapsulate logic.
Parameters and Return Values: The chapters explore how functions can take parameters and return values, reinforcing the idea of data flow within programs. In class, we practiced writing functions that processed user input and returned results, demonstrating how these concepts enable dynamic and flexible programming.
Lists and Mutability: Chapter 8 introduces lists as a way to store collections of data, highlighting their mutable nature. We’ve touched on lists in class when discussing how to manage multiple pieces of related data, such as storing grades or inventory items, which aligns with the idea of using lists for organization.
Dictionaries: Chapter 9 focuses on dictionaries as a way to handle key-value pairs. This connects to our classroom work on data structures, as we’ve explored how dictionaries allow for efficient data retrieval and organization, similar to how we used them to map user IDs to user information in our projects.
Iteration and Looping: The concept of iteration is covered in these chapters, particularly in relation to lists and dictionaries. In class, we practiced using loops to iterate over collections, reinforcing the idea of automating repetitive tasks. For example, we wrote loops to calculate averages or to format lists, directly applying the principles discussed in the textbook.

Chapter 11: Debugging
Content: Focuses on techniques to identify and fix errors in code, including syntax errors, runtime errors, and logical errors.
Vocabulary: Debugging, errors, exceptions, print statement, assertions.
Chapter 12: Data Structures
Content: Introduces lists and tuples as fundamental data structures for storing collections of data.
Vocabulary: List, tuple, indexing, slicing, mutation.
Chapter 13: Case Study: Search
Content: Explores algorithms for searching through data, specifically linear search and binary search.
Vocabulary: Search, linear search, binary search, algorithm, efficiency.
Chapter 15: Lists
Content: Dives deeper into list operations and methods, including list comprehensions and advanced list manipulations.
Vocabulary: List comprehensions, append, extend, sort, slice.
Chapter 16: Dictionaries
Content: Introduces dictionaries as a way to store key-value pairs and discusses their usage and methods.
Vocabulary: Dictionary, key, value, items, keys, values.
Chapter 17: Case Study: Data Structures
Content: Applies knowledge of data structures to solve problems, reinforcing concepts learned in previous chapters.
Vocabulary: Data structure, stack, queue, implementation, abstraction.


# Question 2: Understanding loops

### Part A
 For these fragments of code, write the output **_without_** running them.

```python
    num1 = 10
    while num1 > 3:
        print(num1)
        num1 = num1 - 1
```

 #### Output:
10
9
8
7
6
5
4

```python
    num2 = 10
    while True:
        if num2 < 7:
            # break makes us quit the loop
            break
        print(num2)
        num2 = num2 - 1
```

10
9
8
7


```python
    divisor = 2
    for i in range(0, 10, 2):
        print(i / divisor)
```

0.0
1.0
2.0
3.0
4.0


### Part B

These programs use `for` loops to iterate through strings.

```python
    my_string = 'stressed'
    for letter in my_string:
        print(letter)
```

s
t
r
e
s
s
e
d


```python
    my_string = 'stressed'
    res = ''
    for letter in my_string:
        res = letter + res
    print(res)
```

desserts


Describe, in your own words, what’s happening in this loop:

(edit this cell with your answer)

# Question 3 : Exponentials

Imagine that the `**` operator for numbers has now been removed from the Python programming language. Write a program that prompts the user for a base and an exponent and uses a `for` loop to calculate `base ** exp`. _**Note:** You should handle the case where the user does not type a positive integer._

```python
def exponential(base, exp):
    if exp < 0:
        result = 1
        for _ in range(-exp):
            result /= base
    else:
        result = 1
        for _ in range(exp):
            result *= base
    return result

```

```python
base = float(input("Enter a base: "))
exp = int(input("Enter an exponent (positive integer): "))

if exp < 0:
    print("Please enter a positive integer for the exponent.")
else:
    print(f"The result of {base} raised to the power of {exp} is {exponential(base, exp)}")

```

Output:
```
Enter a base: 2
Enter an exponent (positive integer): 4
The result of 2.0 raised to the power of 4 is 16.0
```

The cells below are autograder cells; don't edit those! But you can run them (using the Run button) to test your own code! If they don't spit out any error messages, then your code works!

```python
assert exponential(2, 4) == 16
assert exponential(3, 0) == 1
assert exponential(-2, 3) == -8
assert exponential(2, -2) == 0.25


```

# Question 4 : Functions and more functions

### Part A

Write a function:

 - named `list_of_positive_indices`
 - takes 1 input (i.e., argument), a list
 - returns 1 value, a list

In this function, you should loop through each element of the argument list. If the element is *positive* (i.e., greater than zero), then you should append the *index number of the positive element* to a new list. If not, skip that index (i.e. do nothing). Once the loop is finished, return the new list.

For example, `list_of_positive_indices([1, -1, 0, 3])` should return `[0, 3]` because these are the **indices** of 1 and 3. No imports!

*HINT*: Instead of doing a simple `for element in list:` loop, try using the `range()` function in the loop header. What will you be looping through instead of the elements themselves?

```python
def list_of_positive_indices(in_list):
    indices = []
    for i in range(len(in_list)):
        if in_list[i] > 0:
            indices.append(i)
    return indices

```

Run the autograder cells below to test your code!

<!-- grade_id: q1c_test1, points: 5, grade: true, solution: false -->
```python
def is_even(num):
    return num % 2 == 0

```

<!-- grade_id: q1c_test2, points: 5, grade: true, solution: false -->
```python

# Define num before using it
num = 5  # Or any other desired value

is_odd = not is_even(num)
```

### Part B

Write a function that tests if a number is even.

*HINT*: Recall the modulo operator `%`. This performs division on two numbers, but instead of returning their quotient, it returns the *remainder* of the division. For example, `7 % 3` returns 1, since the quotient is 2 with a remainder of 1.

Your function should:

 - be named `is_even`
 - take 1 argument, a number
 - return `True` if the argument is even, `False` otherwise

Make sure you're returning a *boolean*!

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def is_even(num):
    pass
```

<!-- grade_id: q2a_test1, points: 2.5, grade: true, solution: false -->
```python
assert is_even(3) == False
```

<!-- grade_id: q2a_test2, points: 2.5, grade: true, solution: false -->
```python
assert is_even(123456789) == False
```

<!-- grade_id: q2a_test3, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42) == True
```

<!-- grade_id: q2a_test4, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42.1) == False
```

### Part C

How could you test if a number is *odd*, using your `is_even` function from Part B?

Write your answer, either using 1 line of code or explaining in only 1 sentence, in the box below (double click on the statement below to activate the box).

<!-- grade_id: q2b, points: 5, grade: true, solution: true -->
is_odd = not is_even(num)


### Part D *(Optional)*

*This part is optional*. Write a function that finds the *second-smallest* number in a list of numbers. Your function should:

 - be named `second_smallest`
 - take 1 argument, a list of numbers
 - return 1 number, the second-smallest number in the list

Take note of that last point. There's a possibility, after all, that the list your function receives as input will have 1 or even 0 elements in it. How, then, do you return the second-smallest element of a 0-element list? Well, simply, you can't. So instead, you'll return False.

For example, `second_smallest([1, 2, 3])` should return `2`. Another example `second_smallest([10, 100, 1000, 10000])` should return 100. And `second_smallest([1])` should return `False`.

You **cannot** use any built-in sorting functions! Only `range()` and `len()`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
def second_smallest(in_list):
    pass
```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(89547)

l1 = np.random.randint(-100, 100, 100)
a1 = np.sort(l1)[1]
np.testing.assert_allclose(a1, second_smallest(l1.tolist()))
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(485)

l2 = np.random.randint(0, 1000, 100)
a2 = np.sort(l2)[1]
np.testing.assert_allclose(a2, second_smallest(l2.tolist()))
```

<!-- grade_id: q2c_test3, points: 2.5, grade: true, solution: false -->
```python
assert second_smallest([]) == False
```

<!-- grade_id: q2c_test4, points: 2.5, grade: true, solution: false -->
```python
assert second_smallest([100]) == False
```

# Question 5: Greetings

### Part A

Print a greeting saying “Hello, Harry” and “Goodbye, Potter” using only the variables provided below. Complete instructions below.

```python
# Greeting
print('*******************')
print('*                 *')
print('* Hello, Harry!!! *')
print('*                 *')
print('*******************')

# Goodbye
print('***********************')
print('*                     *')
print('* Goodbye, Potter!!!  *')
print('*                     *')
print('***********************')

```

Output:
```
*******************
*                 *
* Hello, Harry!!! *
*                 *
*******************
***********************
*                     *
* Goodbye, Potter!!!  *
*                     *
***********************
```

### Part B

 **Reflection:** The output of `greetings` looks very similar to the one from PS0. However, the code looks different. How would you compare the two versions? What advantages/disadvantages do you see in the two approaches?

More flexibility and consistency are provided by the new welcome output technique, which has an organized style based exclusively on predefined variables. The previous version was less versatile, despite being considerably easier to understand. Updates and alterations to the content can be easily done without changing the overall structure by using variables.

# Question 6: Acronym

This exercise will allow you to practice with the methods for strings and lists. An acronym is a word formed by taking the first letters of the words in a phrase and making a word from them. For example, RAM is an acronym for “random access memory.” In this exercise, you are going to write a program that should allow the user to type in a phrase, and the computer should then output the acronym for that phrase. An example of how your program should work is as follows:

**Note:** The acronym should be all uppercase, even if the words in the phrase are not capitalized.

Enter a phrase: loyola marymount university
Your acronym is: LMU
  

```python
phrase = input("Enter a phrase: ")
words = phrase.split()
acronym = ''.join(word[0].upper() for word in words)
print(f"Your acronym is: {acronym}")

```

Output:
```
Enter a phrase:  loyola marymount university
Your acronym is: LMU
```